// Copyright (C) 2015 the V8 project authors. All rights reserved.
// This code is governed by the BSD license found in the LICENSE file.
/*---
esid: sec-array.prototype.findlastindex
description: >
  Predicate is called for each array property.
info: |
  Array.prototype.findLastIndex ( predicate[ , thisArg ] )

  ...
  4. Let k be len - 1.
  5. Repeat, while k â‰¥ 0,
    ...
    c. Let testResult be ! ToBoolean(? Call(predicate, thisArg, Â« kValue, ğ”½(k), O Â»)).
  ...
---*/

var arr = [undefined, , , 'foo'];
var called = 0;

arr.findLastIndex(function() {
  called++;
});

assert.sameValue(called, 4);
