// Copyright (C) 2015 the V8 project authors. All rights reserved.
// This code is governed by the BSD license found in the LICENSE file.
/*---
esid: sec-array.prototype.findlast
description: >
  Predicate is only called if this.length is > 0.
info: |
  Array.prototype.findLast ( predicate[ , thisArg ] )

  ...
  4. Let k be len - 1.
  5. Repeat, while k â‰¥ 0,
    ...
    c. Let testResult be ! ToBoolean(? Call(predicate, thisArg, Â« kValue, ğ”½(k), O Â»)).
  ...
  6. Return undefined.
---*/

var called = false;

var predicate = function() {
  called = true;
  return true;
};

var result = [].findLast(predicate);

assert.sameValue(called, false, '[].findLast(predicate) does not call predicate');
assert.sameValue(result, undefined, '[].findLast(predicate) returned undefined');
