HARNESS_FEATURES - Missing from `features`: Symbol.toPrimitive
^ expected errors | v input
// Copyright (C) 2017 Rick Waldron. All rights reserved.
// This code is governed by the BSD license found in the LICENSE file.
/*---
esid: sec-assignment-operators-static-semantics-early-errors
description: Minimal test
features: [BigInt]
includes: [typeCoercion.js]
---*/

// empty
