// Copyright (C) 2026 Michael Ficarra. All rights reserved.
// This code is governed by the BSD license found in the LICENSE file.
/*---
esid: sec-iterator.prototype.chunks
description: >
  Iterator.prototype.chunks is a built-in function
features: [iterator-chunking]
---*/

assert.sameValue(typeof Iterator.prototype.chunks, 'function');
