// Copyright (C) 2026 Michael Ficarra. All rights reserved.
// This code is governed by the BSD license found in the LICENSE file.
/*---
esid: sec-iterator.prototype.chunks
description: >
  Underlying iterator next returns object with throwing done getter
info: |
  Iterator.prototype.chunks ( chunkSize )

features: [iterator-chunking]
---*/
class ThrowingIterator extends Iterator {
  next() {
    return {
      get done() {
        throw new Test262Error();
      },
      value: 1,
    };
  }
  return() {
    throw new Error();
  }
}

let iterator = new ThrowingIterator().chunks(1);

assert.throws(Test262Error, function () {
  iterator.next();
});
