// Copyright (C) 2024 Igalia, S. L. All rights reserved.
// This code is governed by the BSD license found in the LICENSE file.

/*---
esid: sec-asynccontext-variable-constructor
description: >
  Property descriptor of AsyncContext.Variable
info: |
  17 ECMAScript Standard Built-in Objects:

  Every other data property described in clauses 18 through 26 and in Annex B.2
  has the attributes { [[Writable]]: true, [[Enumerable]]: false,
  [[Configurable]]: true } unless otherwise specified.
includes: [propertyHelper.js]
features: [AsyncContext]
---*/

verifyProperty(AsyncContext, 'Variable', {
  enumerable: false,
  writable: true,
  configurable: true
});
