// Copyright (C) 2026 Michael Ficarra. All rights reserved.
// This code is governed by the BSD license found in the LICENSE file.
/*---
esid: sec-iterator.prototype.windows
description: >
  Iterator.prototype.windows is a built-in function
features: [iterator-chunking]
---*/

assert.sameValue(typeof Iterator.prototype.windows, 'function');
