// Copyright (C) 2026 Igalia, S.L. All rights reserved.
// This code is governed by the BSD license found in the LICENSE file.

/*---
esid: sec-temporal.plaindate.from
description: >
  The ISO date of Nowruz (1 Farvardin) for Persian years 1206–1498 AP matches
  the data published by the Iranian calendar authority
info: |
  The Iranian calendar authority has published the Gregorian date of 1 Farvardin
  for each Persian year from 1206 to 1498 AP.
  Source: https://calendar.ut.ac.ir/documents/2139738/7092644/Kabise+Shamsi+1206-1498.pdf
  Data: https://github.com/roozbehp/persiancalendar/blob/main/kabise.txt (CC0)
features: [Temporal, Intl.Era-monthcode]
---*/

// [persianYear, isoDate] pairs from the Iranian calendar authority
const nowruzDates = [
  [1206, "1827-03-22"],
  [1207, "1828-03-21"],
  [1208, "1829-03-21"],
  [1209, "1830-03-21"],
  [1210, "1831-03-21"],
  [1211, "1832-03-21"],
  [1212, "1833-03-21"],
  [1213, "1834-03-21"],
  [1214, "1835-03-21"],
  [1215, "1836-03-21"],
  [1216, "1837-03-21"],
  [1217, "1838-03-21"],
  [1218, "1839-03-21"],
  [1219, "1840-03-21"],
  [1220, "1841-03-21"],
  [1221, "1842-03-21"],
  [1222, "1843-03-21"],
  [1223, "1844-03-21"],
  [1224, "1845-03-21"],
  [1225, "1846-03-21"],
  [1226, "1847-03-21"],
  [1227, "1848-03-21"],
  [1228, "1849-03-21"],
  [1229, "1850-03-21"],
  [1230, "1851-03-21"],
  [1231, "1852-03-21"],
  [1232, "1853-03-21"],
  [1233, "1854-03-21"],
  [1234, "1855-03-21"],
  [1235, "1856-03-21"],
  [1236, "1857-03-21"],
  [1237, "1858-03-21"],
  [1238, "1859-03-21"],
  [1239, "1860-03-21"],
  [1240, "1861-03-21"],
  [1241, "1862-03-21"],
  [1242, "1863-03-21"],
  [1243, "1864-03-20"],
  [1244, "1865-03-21"],
  [1245, "1866-03-21"],
  [1246, "1867-03-21"],
  [1247, "1868-03-20"],
  [1248, "1869-03-21"],
  [1249, "1870-03-21"],
  [1250, "1871-03-21"],
  [1251, "1872-03-20"],
  [1252, "1873-03-21"],
  [1253, "1874-03-21"],
  [1254, "1875-03-21"],
  [1255, "1876-03-20"],
  [1256, "1877-03-21"],
  [1257, "1878-03-21"],
  [1258, "1879-03-21"],
  [1259, "1880-03-20"],
  [1260, "1881-03-21"],
  [1261, "1882-03-21"],
  [1262, "1883-03-21"],
  [1263, "1884-03-20"],
  [1264, "1885-03-21"],
  [1265, "1886-03-21"],
  [1266, "1887-03-21"],
  [1267, "1888-03-20"],
  [1268, "1889-03-21"],
  [1269, "1890-03-21"],
  [1270, "1891-03-21"],
  [1271, "1892-03-20"],
  [1272, "1893-03-21"],
  [1273, "1894-03-21"],
  [1274, "1895-03-21"],
  [1275, "1896-03-20"],
  [1276, "1897-03-20"],
  [1277, "1898-03-21"],
  [1278, "1899-03-21"],
  [1279, "1900-03-21"],
  [1280, "1901-03-21"],
  [1281, "1902-03-22"],
  [1282, "1903-03-22"],
  [1283, "1904-03-21"],
  [1284, "1905-03-21"],
  [1285, "1906-03-22"],
  [1286, "1907-03-22"],
  [1287, "1908-03-21"],
  [1288, "1909-03-21"],
  [1289, "1910-03-22"],
  [1290, "1911-03-22"],
  [1291, "1912-03-21"],
  [1292, "1913-03-21"],
  [1293, "1914-03-22"],
  [1294, "1915-03-22"],
  [1295, "1916-03-21"],
  [1296, "1917-03-21"],
  [1297, "1918-03-22"],
  [1298, "1919-03-22"],
  [1299, "1920-03-21"],
  [1300, "1921-03-21"],
  [1301, "1922-03-22"],
  [1302, "1923-03-22"],
  [1303, "1924-03-21"],
  [1304, "1925-03-21"],
  [1305, "1926-03-22"],
  [1306, "1927-03-22"],
  [1307, "1928-03-21"],
  [1308, "1929-03-21"],
  [1309, "1930-03-21"],
  [1310, "1931-03-22"],
  [1311, "1932-03-21"],
  [1312, "1933-03-21"],
  [1313, "1934-03-21"],
  [1314, "1935-03-22"],
  [1315, "1936-03-21"],
  [1316, "1937-03-21"],
  [1317, "1938-03-21"],
  [1318, "1939-03-22"],
  [1319, "1940-03-21"],
  [1320, "1941-03-21"],
  [1321, "1942-03-21"],
  [1322, "1943-03-22"],
  [1323, "1944-03-21"],
  [1324, "1945-03-21"],
  [1325, "1946-03-21"],
  [1326, "1947-03-22"],
  [1327, "1948-03-21"],
  [1328, "1949-03-21"],
  [1329, "1950-03-21"],
  [1330, "1951-03-22"],
  [1331, "1952-03-21"],
  [1332, "1953-03-21"],
  [1333, "1954-03-21"],
  [1334, "1955-03-22"],
  [1335, "1956-03-21"],
  [1336, "1957-03-21"],
  [1337, "1958-03-21"],
  [1338, "1959-03-22"],
  [1339, "1960-03-21"],
  [1340, "1961-03-21"],
  [1341, "1962-03-21"],
  [1342, "1963-03-21"],
  [1343, "1964-03-21"],
  [1344, "1965-03-21"],
  [1345, "1966-03-21"],
  [1346, "1967-03-21"],
  [1347, "1968-03-21"],
  [1348, "1969-03-21"],
  [1349, "1970-03-21"],
  [1350, "1971-03-21"],
  [1351, "1972-03-21"],
  [1352, "1973-03-21"],
  [1353, "1974-03-21"],
  [1354, "1975-03-21"],
  [1355, "1976-03-21"],
  [1356, "1977-03-21"],
  [1357, "1978-03-21"],
  [1358, "1979-03-21"],
  [1359, "1980-03-21"],
  [1360, "1981-03-21"],
  [1361, "1982-03-21"],
  [1362, "1983-03-21"],
  [1363, "1984-03-21"],
  [1364, "1985-03-21"],
  [1365, "1986-03-21"],
  [1366, "1987-03-21"],
  [1367, "1988-03-21"],
  [1368, "1989-03-21"],
  [1369, "1990-03-21"],
  [1370, "1991-03-21"],
  [1371, "1992-03-21"],
  [1372, "1993-03-21"],
  [1373, "1994-03-21"],
  [1374, "1995-03-21"],
  [1375, "1996-03-20"],
  [1376, "1997-03-21"],
  [1377, "1998-03-21"],
  [1378, "1999-03-21"],
  [1379, "2000-03-20"],
  [1380, "2001-03-21"],
  [1381, "2002-03-21"],
  [1382, "2003-03-21"],
  [1383, "2004-03-20"],
  [1384, "2005-03-21"],
  [1385, "2006-03-21"],
  [1386, "2007-03-21"],
  [1387, "2008-03-20"],
  [1388, "2009-03-21"],
  [1389, "2010-03-21"],
  [1390, "2011-03-21"],
  [1391, "2012-03-20"],
  [1392, "2013-03-21"],
  [1393, "2014-03-21"],
  [1394, "2015-03-21"],
  [1395, "2016-03-20"],
  [1396, "2017-03-21"],
  [1397, "2018-03-21"],
  [1398, "2019-03-21"],
  [1399, "2020-03-20"],
  [1400, "2021-03-21"],
  [1401, "2022-03-21"],
  [1402, "2023-03-21"],
  [1403, "2024-03-20"],
  [1404, "2025-03-21"],
  [1405, "2026-03-21"],
  [1406, "2027-03-21"],
  [1407, "2028-03-20"],
  [1408, "2029-03-20"],
  [1409, "2030-03-21"],
  [1410, "2031-03-21"],
  [1411, "2032-03-20"],
  [1412, "2033-03-20"],
  [1413, "2034-03-21"],
  [1414, "2035-03-21"],
  [1415, "2036-03-20"],
  [1416, "2037-03-20"],
  [1417, "2038-03-21"],
  [1418, "2039-03-21"],
  [1419, "2040-03-20"],
  [1420, "2041-03-20"],
  [1421, "2042-03-21"],
  [1422, "2043-03-21"],
  [1423, "2044-03-20"],
  [1424, "2045-03-20"],
  [1425, "2046-03-21"],
  [1426, "2047-03-21"],
  [1427, "2048-03-20"],
  [1428, "2049-03-20"],
  [1429, "2050-03-21"],
  [1430, "2051-03-21"],
  [1431, "2052-03-20"],
  [1432, "2053-03-20"],
  [1433, "2054-03-21"],
  [1434, "2055-03-21"],
  [1435, "2056-03-20"],
  [1436, "2057-03-20"],
  [1437, "2058-03-21"],
  [1438, "2059-03-21"],
  [1439, "2060-03-20"],
  [1440, "2061-03-20"],
  [1441, "2062-03-20"],
  [1442, "2063-03-21"],
  [1443, "2064-03-20"],
  [1444, "2065-03-20"],
  [1445, "2066-03-20"],
  [1446, "2067-03-21"],
  [1447, "2068-03-20"],
  [1448, "2069-03-20"],
  [1449, "2070-03-20"],
  [1450, "2071-03-21"],
  [1451, "2072-03-20"],
  [1452, "2073-03-20"],
  [1453, "2074-03-20"],
  [1454, "2075-03-21"],
  [1455, "2076-03-20"],
  [1456, "2077-03-20"],
  [1457, "2078-03-20"],
  [1458, "2079-03-21"],
  [1459, "2080-03-20"],
  [1460, "2081-03-20"],
  [1461, "2082-03-20"],
  [1462, "2083-03-21"],
  [1463, "2084-03-20"],
  [1464, "2085-03-20"],
  [1465, "2086-03-20"],
  [1466, "2087-03-21"],
  [1467, "2088-03-20"],
  [1468, "2089-03-20"],
  [1469, "2090-03-20"],
  [1470, "2091-03-21"],
  [1471, "2092-03-20"],
  [1472, "2093-03-20"],
  [1473, "2094-03-20"],
  [1474, "2095-03-20"],
  [1475, "2096-03-20"],
  [1476, "2097-03-20"],
  [1477, "2098-03-20"],
  [1478, "2099-03-20"],
  [1479, "2100-03-21"],
  [1480, "2101-03-21"],
  [1481, "2102-03-21"],
  [1482, "2103-03-21"],
  [1483, "2104-03-21"],
  [1484, "2105-03-21"],
  [1485, "2106-03-21"],
  [1486, "2107-03-21"],
  [1487, "2108-03-21"],
  [1488, "2109-03-21"],
  [1489, "2110-03-21"],
  [1490, "2111-03-21"],
  [1491, "2112-03-21"],
  [1492, "2113-03-21"],
  [1493, "2114-03-21"],
  [1494, "2115-03-21"],
  [1495, "2116-03-21"],
  [1496, "2117-03-21"],
  [1497, "2118-03-21"],
  [1498, "2119-03-21"],
];

for (const [persianYear, expectedISO] of nowruzDates) {
  const date = Temporal.PlainDate.from({
    year: persianYear,
    month: 1,
    day: 1,
    calendar: "persian"
  });
  const isoDate = date.withCalendar("iso8601");
  assert.sameValue(
    isoDate.toString(),
    expectedISO,
    "Nowruz (1 Farvardin) of Persian year " + persianYear + " should be ISO " + expectedISO
  );
}
