// Copyright (C) 2020 Alexey Shvayka. All rights reserved.
// This code is governed by the BSD license found in the LICENSE file.

/*---
esid: sec-names-and-keywords
description: Test astral characters of ID_Continue.
info: |
  Astral characters (Other_ID_Continue)
---*/

// Other_ID_Continue (Unicode 4.1)
var ağ’¤; // U+104A4
var ağ¨; // U+10A0F

// Other_ID_Continue (Unicode 5.2)
var ağ‘‚º; // U+110BA
var ağ‘‚¸; // U+110B8

// Other_ID_Continue (Unicode 6.3)
var ağ‘„; // U+11101
var ağ‘ª; // U+1106A

// Other_ID_Continue (Unicode 7.0)
var ağ‘–¿; // U+115BF
var ağ‘Œ; // U+11301

// Other_ID_Continue (Unicode 8.0)
var ağ¨…; // U+1DA05
var ağ‘œª; // U+1172A

// Other_ID_Continue (Unicode 9.0)
var ağ‘±—; // U+11C57
var ağ‘²Ÿ; // U+11C9F

// Other_ID_Continue (Unicode 10.0)
var ağ‘¨¹; // U+11A39
var ağ‘µ; // U+11D41

// Other_ID_Continue (Unicode 11.0)
var ağ´±; // U+10D31
var ağ½; // U+10F4D

// Other_ID_Continue (Unicode 12.0)
var ağ„µ; // U+1E135
var ağ…ƒ; // U+1E143
