// Copyright (C) 2020 Alexey Shvayka. All rights reserved.
// This code is governed by the BSD license found in the LICENSE file.

/*---
esid: sec-names-and-keywords
description: Test astral characters of ID_Start.
info: |
  Astral characters (Other_ID_Start)
---*/

// Other_ID_Start (Unicode 4.1)
var ğ¨ª; // U+10A2A
var ğ…«; // U+1016B

// Other_ID_Start (Unicode 5.2)
var ğ“¬; // U+1306C
var ğ’…¹; // U+12179

// Other_ID_Start (Unicode 6.3)
var ğ–¼Ÿ; // U+16F1F
var ğ‘„¥; // U+11125

// Other_ID_Start (Unicode 7.0)
var ğ „; // U+1E804
var ğ¢—; // U+1E897

// Other_ID_Start (Unicode 8.0)
var ğ’“‡; // U+124C7
var ğ³¢; // U+10CE2

// Other_ID_Start (Unicode 9.0)
var ğ˜‹; // U+182C1
var ğ—Œ¹; // U+17339

// Other_ID_Start (Unicode 10.0)
var ğ›‹; // U+1B2D0
var ğ›‚‘; // U+1B091

// Other_ID_Start (Unicode 11.0)
var ğ˜Ÿ®; // U+187EE
var ğ´; // U+10D01

// Other_ID_Start (Unicode 12.0)
var ğ‹§; // U+1E2E7
var ğ–½†; // U+16F46
