INCLUDES
^ expected errors | v input
// Copyright (C) 2026 Igalia, S.L. All rights reserved.
// This code is governed by the BSD license found in the LICENSE file.
/*---
esid: sec-foobar
description: >
  Includes a harness file that doesn't have any defines (e.g. deprecated
  compareArray.js)
includes: [noDefines.js]
---*/

void 0;
