// Copyright 2009 the Sputnik authors.  All rights reserved.
// This code is governed by the BSD license found in the LICENSE file.

/*---
esid: sec-future-reserved-words
description: Checking if execution of "abstract=1" succeeds
info: |
  `abstract` was removed from the list of future reserved words in ECMAScript 5.
---*/

var abstract = 1;
