// Copyright 2009 the Sputnik authors.  All rights reserved.
// This code is governed by the BSD license found in the LICENSE file.

/*---
esid: sec-future-reserved-words
description: Checking if execution of "byte=1" succeeds
info: |
  `byte` was removed from the list of future reserved words in ECMAScript 5.
---*/

var byte = 1;
